# Statystyka matematyczna - wiczenia laboratoryjne 2023/2024

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu.
- Zadanie nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj za zadanie wyapa podstawowe bdy. Przejcie wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu. 
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem).
- Rozmiar poziomy figur nie powinien przekracza 20 cali.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (nazwa zespou oraz imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespou: 2.38 Czonkowie: Roksana Jandura , Katarzyna Wesowoska , Katarzyna Tokarczuk , Magdalena Pogorzelec 

---

# Zestaw zada 8: Estymacja metod bootstrap

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakiet贸w wchodzcych w skad standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakiet贸w
```

### Dane do zada

W celu wygenerowania danych wykorzystywanych w zawartych w notatniku zadaniach i kom贸rkach testowych wykonaj obie poni偶sze kom贸rki.

```python
sample_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
sample_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=25, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=25, random_state=13)
sample_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["X", "Y"])
sample_2.head()
```

Output:
```
X         Y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

### Zadanie 1: Estymacja bootstapowa parametr贸w pr贸by losowej [8 pkt]

Przygotuj funkcj `bootstrap_parameter_estimation()`, kt贸ra korzystajc z metody bootstrap bdzie dokonywaa estymacji przedziaowej podstawowych parametr贸w pr贸by losowej - wartoci oczekiwanej, wariancji i odchylenia standardowego dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

W obliczeniach skorzystaj z nieobci偶onych estymator贸w wariancji i odchylenia standardowego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`P nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie P (symbol lub nazwa estymowanego parametru), X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakadaj losowanie pr贸b wg nastpujcego schematu - kolejne pr贸by s losowane w pli for, a funkcja losujca pr贸b (np. `pd.sample()`) w kolejnych losowaniach ma ustawion warto parametru odpowiadajcego za ziarno generatora liczb losowych na warto `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwizaniu nie nale偶y korzysta z gotowych funkcji su偶cych do estymacji bootstrapowej.

<!-- grade_id: cell-618e12c8ffef0965, points: 0, grade: false, solution: true -->
```python
def bootstrap_parameter_estimation(data, parameter, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej.
    parameter: str
        W zale偶noci od estymowanego parametru przyjmuje warto "mean", "var" lub "std".
    number_of_samples: int
        Liczba pr贸b bootstrapowych.
    alpha: float
        Warto okrelajca poziom istotnoci.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    parameter_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    parameter_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """       
    estimates = [] #lista z mean,var,std dla kazdej pr贸by bo0otstrap
    #ptla d losowania kolejnych pr贸b
    for i in range(number_of_samples):
        #ziarno generatora liczb losowych zwieksza sie za kazda iteracj0
        sample = data.sample(n=len(data), replace=True, random_state=random_state + i)
        
        if parameter == "mean":
            est = sample['X'].mean()
        elif parameter == "var":
            est = sample['X'].var(ddof=1) #ddof=1 bo nieobci偶ony estymator wariancji
        elif parameter == "std":
            est = sample['X'].std(ddof=1) #ddof=1 bo nieobci偶ony estymator std
        estimates.append(est) #dodaje  obliczona wartosc do listy
    
    lower_percentile = 100 * (alpha / 2) #dolny procentyl
    upper_percentile = 100 * (1 - alpha / 2) #g贸rny percentyl
    parameter_low = np.percentile(estimates, lower_percentile) #dolna granica
    parameter_high = np.percentile(estimates, upper_percentile) #g贸rna granica
    
    Z = 1 - alpha #poziom ufnoci
    print(f"{parameter} nale偶y do przedziau [{parameter_low:.4f}, {parameter_high:.4f}] przy zao偶eniu poziomu ufnoci 1-伪 = {Z}")
    
    return parameter_low, parameter_high

```

<!-- grade_id: cell-eab0d556900c97d4, points: 8, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "mean", 1000, 0.05, 34), (4.913578501079754, 5.058598650238277)))
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "var", 2000, 0.02, 31), (0.009780029754528753, 0.05012098474121563)))
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "std", 1500, 0.1, 39), (0.11848597559111698, 0.20690074715252046)))
```

Output:
```
mean nale偶y do przedziau [4.9136, 5.0586] przy zao偶eniu poziomu ufnoci 1-伪 = 0.95
var nale偶y do przedziau [0.0098, 0.0501] przy zao偶eniu poziomu ufnoci 1-伪 = 0.98
std nale偶y do przedziau [0.1185, 0.2069] przy zao偶eniu poziomu ufnoci 1-伪 = 0.9
```

### Zadanie 2: Estymacja bootstrapowa wartoci wsp贸czynnika korelacji liniowej [4 pkt]

Przygotuj funkcj `bootstrap_correlation_estimation()`, kt贸ra korzystajc z metody bootstrap bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnika korelacji liniowej Pearsona dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`r nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakadaj losowanie pr贸b wg nastpujcego schematu - kolejne pr贸by s losowane w pli for, a funkcja losujca pr贸b (np. `pd.sample()`) w kolejnych losowaniach ma ustawion warto parametru odpowiadajcego za ziarno generatora liczb losowych na warto `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwizaniu nie nale偶y korzysta z gotowych funkcji su偶cych do estymacji bootstrapowej.

<!-- grade_id: cell-d785d7ff406d589d, points: 0, grade: false, solution: true -->
```python
def bootstrap_correlation_estimation(data, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    number_of_samples: int
        Liczba pr贸b bootstrapowych.
    alpha: float
        Warto okrelajca poziom istotnoci.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    r_corr_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    # YOUR CODE HERE
    x = data['X']
    y = data['Y']

    # wsp贸czynniki korelacji z pr贸b bootstrapowych
    bootstrap_correlations = []

    for i in range(number_of_samples):
        bootstrap_sample = data.sample(n=len(data), replace=True, random_state=random_state + i)
        correlation, _ = sp.stats.pearsonr(bootstrap_sample['X'], bootstrap_sample['Y'])
        bootstrap_correlations.append(correlation)

    bootstrap_correlations = sorted(bootstrap_correlations)

    lower_index = 100 * (alpha / 2)
    upper_index = 100 * (1 - alpha / 2)

    r_corr_low = np.percentile(bootstrap_correlations, lower_index)
    r_corr_high =  np.percentile(bootstrap_correlations, upper_index)
    
    confidence_level = 1 - alpha

    print(f"r nale偶y do przedziau [{r_corr_low:.4f}, {r_corr_high:.4f}] przy zao偶eniu poziomu ufnoci 1- = {confidence_level:.2f}")

    return r_corr_low, r_corr_high
```

<!-- grade_id: cell-4dfca86f32f1359a, points: 4, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(bootstrap_correlation_estimation(sample_2, 1200, 0.05, 54), (0.9292256084679922, 0.9816422271248535)))
assert np.all(np.isclose(bootstrap_correlation_estimation(sample_2, 1500, 0.02, 14), (0.9171033819425604, 0.9840550768251518)))
```

Output:
```
r nale偶y do przedziau [0.9292, 0.9816] przy zao偶eniu poziomu ufnoci 1- = 0.95
r nale偶y do przedziau [0.9171, 0.9841] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

### Zadanie 3: Estymacja bootstrapowa wsp贸czynnik贸w r贸wnania regresji liniowej [6 pkt]

Przygotuj funkcj `bootstrap_linear_regression_coefficients_estimation()`, kt贸ra korzystajc z metody bootstrap bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnik贸w r贸wnania regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`a nale偶y do przedziau [X1, Y1] przy zao偶eniu poziomu ufnoci 1- = Z1` </br>
`b nale偶y do przedziau [X2, Y2] przy zao偶eniu poziomu ufnoci 1- = Z2`

gdzie X1, X2, Y1, Y2, Z1 i Z2 s automatycznie uzupeniane przez funkcj, a X1, X2, Y1 i Y2 dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakadaj losowanie pr贸b wg nastpujcego schematu - kolejne pr贸by s losowane w pli for, a funkcja losujca pr贸b (np. `pd.sample()`) w kolejnych losowaniach ma ustawion warto parametru odpowiadajcego za ziarno generatora liczb losowych na warto `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwizaniu nie nale偶y korzysta z gotowych funkcji su偶cych do estymacji bootstrapowej.

<!-- grade_id: cell-03ff8749dfe56b09, points: 0, grade: false, solution: true -->
```python
def bootstrap_linear_regression_coefficients_estimation(data, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    number_of_samples: int
        Liczba pr贸b bootstrapowych.
    alpha: float
        Warto okrelajca poziom istotnoci.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    a_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru a (a_low, a_high).
    b_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru b (b_low, b_high).
    """            
    # YOUR CODE HERE
    #raise NotImplementedError()

    a_samples = []
    b_samples = []

    for i in range(number_of_samples):
        # Losuowanie pr贸bki ze zwracaniem
        bootstrap_sample = data.sample(n=len(data), replace=True, random_state=random_state + i)
        
        # rednie X i Y
        mean_x = bootstrap_sample['X'].mean()
        mean_y = bootstrap_sample['Y'].mean()
        
        # wsp贸czynnik kierunkowy a
        licznik = np.sum((bootstrap_sample['X'] - mean_x) * (bootstrap_sample['Y'] - mean_y))
        mianownik = np.sum((bootstrap_sample['X'] - mean_x) ** 2)
        a = licznik / mianownik
        
        # staa b
        b = mean_y - a * mean_x
        
        a_samples.append(a)  # Wsp贸czynnik kierunkowy
        b_samples.append(b)  # Staa (przecicie z osi Y)

    # Obliczanie przedzia贸w ufnoci
    a_low = np.percentile(a_samples, alpha * 100 / 2 )
    a_high = np.percentile(a_samples, 100 * (1 - alpha / 2))
    
    b_low = np.percentile(b_samples, alpha * 100 / 2 )
    b_high = np.percentile(b_samples, 100 * (1 - alpha / 2))

     # Poziom ufnoci
    confidence_level = 1 - alpha
    
    print(f"a nale偶y do przedziau [{a_low:.4f}, {a_high:.4f}] przy zao偶eniu poziomu ufnoci 1- = {confidence_level}")
    print(f"b nale偶y do przedziau [{b_low:.4f}, {b_high:.4f}] przy zao偶eniu poziomu ufnoci 1- = {confidence_level}")
   

    return (a_low, a_high), (b_low, b_high)
```

<!-- grade_id: cell-0febc4ee671ad799, points: 6, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(bootstrap_linear_regression_coefficients_estimation(sample_2, 1800, 0.05, 71), ((1.8361393329086446, 2.2918023790827124), (-5.575502208944547, -3.6854174758043574))))
print()
assert np.all(np.isclose(bootstrap_linear_regression_coefficients_estimation(sample_2, 2100, 0.02, 62), ((1.7964929430496048, 2.337236979144576), (-5.771238137193083, -3.481203748551042))))
```

Output:
```
a nale偶y do przedziau [1.8361, 2.2918] przy zao偶eniu poziomu ufnoci 1- = 0.95
b nale偶y do przedziau [-5.5755, -3.6854] przy zao偶eniu poziomu ufnoci 1- = 0.95

a nale偶y do przedziau [1.7965, 2.3372] przy zao偶eniu poziomu ufnoci 1- = 0.98
b nale偶y do przedziau [-5.7712, -3.4812] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

### Zadanie 4: Estymacja bootstrapowa za pomoc dedykowanych funkcji [7 pkt]

Korzystajc z funkcji `scipy.stats.bootstrap()` dokonaj estymacji przedziaowej nastpujcych parametr贸w na poziomie istotnoci 1-alpha = 0.98:
 - na podstawie zmiennej `sample_1`: rednia, wariancja, odchylenie standardowe,
 - na podstawie zmiannej `sample_2`: wsp贸czynnik korelacji liniowej Pearsona, wsp贸czynniki r贸wnania regresji liniowej a i b.
 
Wyniki zapisz do zmiennych o nazwach `mean_ci`, `var_ci`, `std_ci`, `r_ci`, `a_ci` oraz `b_ci` zawierajcych listy z dolnymi i g贸rnymi granicami przedzia贸w ufnoci. Podczas obliczania wartoci ka偶dego z parametr贸w ustaw warto argumentu funkcji `scipy.stats.bootstrap()` odpowiadajcego za ziarno generatora liczb losowych na 25.

Uwagi do wykonania zadania:
 - Automatyczne testy przewiduj obliczenie wartoci wariancji i odchylenia standardowego dla wartoci parametru `ddof=0`.

<!-- grade_id: cell-5028ed987941ac51, points: 0, grade: false, solution: true -->
```python
# mean
boot_strap = sp.stats.bootstrap((sample_1["X"],), statistic=np.mean, confidence_level=0.98, random_state=25)

mean_ci = [boot_strap.confidence_interval[0], boot_strap.confidence_interval[1]]
mean_ci
```

Output:
```
[4.890098551749027, 5.063500727815335]
```

<!-- grade_id: cell-1b6eb85e5b1da27a, points: 0, grade: false, solution: true -->
```python
# var
boot_strap = sp.stats.bootstrap((sample_1["X"],), statistic=np.var, confidence_level=0.98, random_state=25)

var_ci = [boot_strap.confidence_interval[0], boot_strap.confidence_interval[1]]
var_ci
```

Output:
```
[0.013001391902892368, 0.05670851692597062]
```

<!-- grade_id: cell-7fa6f2c46b14d185, points: 0, grade: false, solution: true -->
```python
# std
boot_strap = sp.stats.bootstrap((sample_1["X"],), statistic=np.std, confidence_level=0.98, random_state=25)

std_ci = [boot_strap.confidence_interval[0], boot_strap.confidence_interval[1]]
std_ci
```

Output:
```
[0.11423521149909441, 0.23829111055240196]
```

<!-- grade_id: cell-185444e7b6b077c2, points: 0, grade: false, solution: true -->
```python
# r
def corr_p(x, y):
    return sp.stats.pearsonr(x, y)[0]

boot_strap = sp.stats.bootstrap((sample_2["X"],sample_2["Y"]), statistic=corr_p, paired=True, confidence_level=0.98, random_state=25)
r_ci = [boot_strap.confidence_interval[0], boot_strap.confidence_interval[1]]
r_ci
```

Output:
```
[0.9156715777126913, 0.9822743695601275]
```

<!-- grade_id: cell-6c7a615fa2793828, points: 0, grade: false, solution: true -->
```python
# a
def reg(x, y):
    regress= sp.stats.linregress(x, y)
    return regress[0], regress[1]
boot_strap = sp.stats.bootstrap((sample_2["X"],sample_2["Y"]), reg, paired=True, confidence_level=0.98, random_state=25)
a_ci = [boot_strap.confidence_interval[0][0], boot_strap.confidence_interval[1][0]]
a_ci
```

Output:
```
[1.792948341896073, 2.3282111601384945]
```

<!-- grade_id: cell-cc1ee5d5033d0923, points: 0, grade: false, solution: true -->
```python
# b
def reg(x, y):
    regress= sp.stats.linregress(x, y)
    return regress[0], regress[1]
boot_strap = sp.stats.bootstrap((sample_2["X"],sample_2["Y"]), sp.stats.linregress, paired=True, confidence_level=0.98, random_state=25)
b_ci = [boot_strap.confidence_interval[0][1], boot_strap.confidence_interval[1][1]]
b_ci
```

Output:
```
[-5.700934784848322, -3.4292283266002834]
```

<!-- grade_id: cell-34b28c6c1e5083f0, points: 7, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(mean_ci, (4.890098551749027, 5.063500727815335)))
assert np.all(np.isclose(var_ci, (0.013001391902892368, 0.05670851692597062)))
assert np.all(np.isclose(std_ci, (0.11423521149909441, 0.23829111055240196)))
assert np.all(np.isclose(r_ci, (0.9156715777126913, 0.9822743695601276)))
assert np.all(np.isclose(a_ci, (1.792948341896073, 2.3282111601384945)))
assert np.all(np.isclose(b_ci, (-5.700934784848322, -3.4292283266002834)))
```

```python

```

```python

```